```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("1_angry_birds_image", "week_9", "lab", assignment_points = 105.0, assignment_tag = 'week9-lab')

# Initialize Otter
import otter
grader = otter.Notebook("1_angry_birds_image.ipynb")
```

# Angry Birds Come Back! 

In a previous assignment we learned how to shoot birds at pigs with unerring accuracy. Now we want to make the game even easier by using image processing to detect scene elements. This is much akin to any computer vision task in science, engineering, manufacturing and medicine. 

In this assignment, we will use image processing to detect the main structure of the scene, and then use that mask to isolate the different elements in the scene. 

We will walk you through the steps of the process. 

Note: You are now getting to be more of a pro at Python. While we do not expect you to know all the syntax off the top of your head, we do expect you to be able to do a little digging and figure out how to do what you need. 

## Importing packages

```python
# We nede to import the following packages:
# - numpy using common alias np
# - matplotlib.pyplot using common alias plt
# - from skimage -- io, color, filters, feature, morphology, measure, segmentation, label, regionprops, clear_border -- note some of these are sub-packages[]

...
```

```python
grader.check("import_packages")
```

```python
# We need to load the image and convert it to grayscale, to load the image assign a variable to the image path, to the name of the image "angry-bird.webp"
...
# load the image using the io.imread function, assign the result to the variable image
...

# Convert to grayscale for edge detection
# use the color.rgb2gray function to convert the image to grayscale, assign the result to the variable gray_image
...

# use the plt.subplots function to create a figure and a single axis, assign the result to the variables fig and ax
...

# use the ax.imshow function to display the image
...
# use the ax.set_title function to set the title of the image - set the title to "Original Image"
...

# use the ax.axis function to turn off the axis labels
...

# use the plt.show function to display the image
...
```

```python
grader.check("load-and-visualize-image")
```

```python
# Apply Canny edge detector to detect edges in the grayscale image
# use the feature.canny function to detect the edges in the image, assign the result to the variable edges
# set the sigma to 2.0
...
# wow I guess I lied when I said we could do this in just a few lines! it is 1 line!

# Visualize the edges
# use the plt.subplots function to create a figure with 2 subplots, assign the result to the variables fig and ax
# the figures should be 10 inches wide and 5 inches tall
# the first subplot should display the grayscale image
# the second subplot should display the edges
# you want the figures layed out in a row with 2 columns
...

# use the ax.imshow function to display the grayscale image in the first subplot -- hint you need to use indexing to get the first subplot
# set the colormap to "gray"
...

# use the ax.set_title function to set the title of the image - set the title to "Grayscale Image"
...

# use the ax.axis function to turn off the axis labels
...
# use the ax.imshow function to display the edges in the second subplot -- hint you need to use indexing to get the second subplot
# set the colormap to "gray"
...

# use the ax.set_title function to set the title of the image - set the title to "Edges Detected (Canny)"
...

# use the ax.axis function to turn off the axis labels
...

# use the plt.show function to display the image
...
```

```python
grader.check("edge-detection")
```

```python
# Dilate the edges to connect broken fragments
# use the morphology.dilation function to dilate the edges, assign the result to the variable dilated_edges
# use the morphology.disk function to create a disk-shaped structuring element with a radius of 3
...

# Fill holes inside the detected edges to create a solid mask
# use the morphology.remove_small_holes function to fill the holes in the edges, assign the result to the variable filled_mask, a threshold of 5000 is used
...

# use the morphology.remove_small_objects function to remove the small objects in the mask, assign the result to the variable filled_mask, a minimum size of 60000 is used
...

# use the plt.subplots function to create a figure with 1 subplot, assign the result to the variables fig and ax
# the figure should be 6 inches wide and 4 inches tall
...

# use the ax.imshow function to display the filled mask
...

# use the ax.set_title function to set the title of the image - set the title to "Filled Mask"
...

# use the ax.axis function to turn off the axis labels
...
# use the plt.show function to display the image
...
```

```python
grader.check("refinement-dilation-and-hole-filling")
```

```python

# use the image.copy function to create a copy of the image, assign the result to the variable structure_only
...

# use the ~ operator to invert the mask, assign the result to the variable inverted_mask
...

# use the structure_only variable to set the background to white, assign the result to the variable structure_only
...

# use the plt.subplots function to create a figure with 1 subplot, assign the result to the variables fig and ax
# the figure should be 6 inches wide and 4 inches tall
...

# use the ax.imshow function to display the structure_only image
...

# use the ax.set_title function to set the title of the image - set the title to "Main Structure Isolated"
...
# use the ax.axis function to turn off the axis labels
...
# use the plt.show function to display the image
...
```

```python
grader.check("isolate-main-structure")
```

```python
# Build the base class for the image processor called ImageProcessor
...
    # initialize the image processor class with the image to analyze
    ...
    
        # assign the image to the image attribute
        ...
        
        # convert the image to HSV color space, assign the result to the hsv_image attribute
        ...
        
        # convert the image to grayscale, assign the result to the gray_image attribute
        ...
        
        # create an empty dictionary to store the masks, labels_dict and regions_dict to their respective attributes
        ...
        
    # define the display_image method, look at the docstring for information on the inputs and outputs
    ...
        """
        Display the image with a given title.

        Parameters:
        title (str): The title of the image to be displayed. Default is "Image".
        """
        ...
```

```python
grader.check("image-processor")
```

## Question 6 (Points: 14.0): Color Mask Class

Now we want to make sure that we can isloate the different objects in the image based on the color. We will use the ColorMask class to help us with this. The color mask class will identify the pixels in the image that fall within a certain range of hue, saturation, and value (brightness) values - and then remove the pixels that are not part of contiguous regions. 


```python
# build the ColorMask class, look at the docstring for information on the inputs and outputs, and the methods
...
    # define the __init__ method, look at the docstring for information on the inputs and outputs
    ...
        """
        Initialize the ColorMask object with HSV image and range values.

        Parameters:
        hsv_img (numpy.ndarray): The HSV image to process.
        hue_range (tuple): The range of hue values to include in the mask.
        sat_range (tuple): The range of saturation values to include in the mask.
        val_range (tuple): The range of value (brightness) values to include in the mask.
        """
        # assign all variables to the respective attributes, add an attribute for the mask initialized to None
        ...

    # define the build_mask method, look at the docstring for information on the inputs and outputs
    ...
        """
        Create a binary mask for pixels whose HSV values fall within the
        specified hue, saturation, and value ranges.

        Returns:
        numpy.ndarray: A binary mask where pixels within the specified ranges are True.
        """
        # extract the hue, saturation, and value channels from the HSV image as H, S, V 
        # The format of the hsv_img is (x,y,3), where x is the height and y is the width, 
        # and the 3 is for the 3 color channels of the image
        ...
        
        # create a binary mask where pixels within the specified ranges are True
        # this is when for example H >= self.hue_range[0] and H <= self.hue_range[1], remember to use the binary & operator
        # you want to ensure all the color conditions are met for the mask to be True
        ...
        
        # return the mask
        return self.mask
    
    # define the clean_mask method, look at the docstring for information on the inputs and outputs
    ...
        """
        Apply morphological operations to clean the binary mask.

        Returns:
        numpy.ndarray: The cleaned binary mask after applying closing and opening operations.
        """
        # apply morphological closing and opening operations to the mask
        # use the morphology.closing function to close the mask, use a disk-shaped structuring element with a radius of 3
        # use the morphology.opening function to open the mask, use a disk-shaped structuring element with a radius of 3
        ...
        
        # return the mask
        ...
```

```python
grader.check("Color-mask-and-merge")
```

## Question 7 (Points: 13.0): Derived Classes for Specific Color Masks

Now we want to create the derived classes for the specific color masks. We will create a class for the pig mask, the wood mask, the gray mask, and the blue mask. This is the power of using classes, one class can be used to create multiple derived classes, and each derived class can have its own unique properties and methods.

```python
# build the PigMask class, look at the docstring for information on the inputs and outputs
...
    """
    A class to represent a mask for detecting pig objects in an image.

    Inherits from the ColorMask base class and initializes the mask with
    specific HSV color range values that correspond to the color of pigs
    in the image.
    """
    # initialize the PigMask class with the hsv_img, hue_range, sat_range, and val_range
    ...
        """
        Constructs all the necessary attributes for the PigMask object.

        Parameters
        ----------
        hsv_img : ndarray
            The HSV image in which the pig mask is to be applied.
        hue_range : tuple, optional
            The range of hue values for the pig mask (default is (0.20, 0.40)).
        sat_range : tuple, optional
            The range of saturation values for the pig mask (default is (0.50, 1.00)).
        val_range : tuple, optional
            The range of value (brightness) values for the pig mask (default is (0.30, 1.00)).
        """
        # initialize the ColorMask class with the hsv_img, hue_range, sat_range, and val_range
        ...


# build the WoodMask class, look at the docstring for information on the inputs and outputs
...
    """
    A class to represent a mask for detecting wood objects in an image.

    Inherits from the ColorMask base class and initializes the mask with
    specific HSV color range values that correspond to the color of wood
    in the image.
    """
    # initialize the WoodMask class with the hsv_img, hue_range, sat_range, and val_range
    ...
        """
        Constructs all the necessary attributes for the WoodMask object.

        Parameters
        ----------
        hsv_img : ndarray
            The HSV image in which the wood mask is to be applied.
        hue_range : tuple, optional
            The range of hue values for the wood mask (default is (0.04, 0.15)).
        sat_range : tuple, optional
            The range of saturation values for the wood mask (default is (0.40, 1.00)).
        val_range : tuple, optional
            The range of value (brightness) values for the wood mask (default is (0.30, 1.00)).
        """
        # initialize the ColorMask class with the hsv_img, hue_range, sat_range, and val_range
        ...


# build the GrayMask class, look at the docstring for information on the inputs and outputs
...
    """
    A class to represent a mask for detecting gray objects in an image.

    Inherits from the ColorMask base class and initializes the mask with
    specific HSV color range values that correspond to the color of gray
    objects in the image.
    """
    # initialize the GrayMask class with the hsv_img, hue_range, sat_range, and val_range
    ...
        """
        Constructs all the necessary attributes for the GrayMask object.

        Parameters
        ----------
        hsv_img : ndarray
            The HSV image in which the gray mask is to be applied.
        hue_range : tuple, optional
            The range of hue values for the gray mask (default is (0.00, 0.833)).
        sat_range : tuple, optional
            The range of saturation values for the gray mask (default is (0.00, 0.52)).
        val_range : tuple, optional
            The range of value (brightness) values for the gray mask (default is (0.30, 0.89)).
        """
        # initialize the ColorMask class with the hsv_img, hue_range, sat_range, and val_range
        ...


# build the BlueMask class, look at the docstring for information on the inputs and outputs
...
    """
    A class to represent a mask for detecting blue objects in an image.

    Inherits from the ColorMask base class and initializes the mask with
    specific HSV color range values that correspond to the color of blue
    objects in the image.
    """
    # initialize the BlueMask class with the hsv_img, hue_range, sat_range, and val_range
    ...
        """
        Constructs all the necessary attributes for the BlueMask object.

        Parameters
        ----------
        hsv_img : ndarray
            The HSV image in which the blue mask is to be applied.
        hue_range : tuple, optional
            The range of hue values for the blue mask (default is (0.52, 0.57)).
        sat_range : tuple, optional
            The range of saturation values for the blue mask (default is (0.141, 0.763)).
        val_range : tuple, optional
            The range of value (brightness) values for the blue mask (default is (0.867, 1.00)).
        """
        # initialize the ColorMask class with the hsv_img, hue_range, sat_range, and val_range
        ...

```

```python
grader.check("PigMask-main")
```

```python
# Now we want to build the MaskBuilder class, look at the docstring for information on the inputs and outputs
...

    # define the __init__ method, look at the docstring for information on the inputs and outputs
    ...
        """
        Initializes the MaskBuilder with an image.

        Parameters
        ----------
        image : ndarray
            The input image to be processed. It is expected to be in a format
            that can be converted to HSV for mask building.
        categories : list of str, optional
            A list of categories for which masks will be built. Default is
            ["pig", "wood", "gray", "blue"].

        Attributes
        ----------
        categories : list of str
            A list of categories for which masks will be built.
        """
        # initialize the ImageProcessor class with the image, as an inheritance
        ...
        # assign the categories to the attributes
        ...

    # define the build_and_clean_masks method, look at the docstring for information on the inputs and outputs
    ...
        """
        Builds and cleans masks for all categories.

        This method creates masks for each category defined in `self.categories`
        using their respective mask classes. It then cleans each mask and stores
        the result in `self.masks`.

        The masks are built and cleaned in the order they appear in `self.categories`.

        Each mask is stored in `self.masks` with the category name as the key.
        If a mask cannot be built, `None` is stored for that category.
        """
        # iterate over each category in self.categories assign to the local variable category
        ...
           
            # get the mask class from the globals
            # this is advanced python, we are using the globals() function to get the mask class from the globals
            # the f-string is used to format the category name to the correct mask class name
            # for example, if the category is "pig", the mask class will be "PigMask"
            mask_class = globals().get(f"{category.capitalize()}Mask")
            
            # if the mask class exists, create an instance of the mask class
            ...
            # assign the mask instance to the masks dictionary with the category name as the key
            # if the mask_instance.build_mask() is not None, then we can clean the mask, and save the output to the masks dictionary
            # otherwise, we save None to the masks dictionary
            ...

    # define the label_and_extract_regions method, look at the docstring for information on the inputs and outputs
    ...
        """
        Labels the cleaned masks and extracts region properties.

        This method iterates over each mask in `self.masks`, labels the mask,
        and extracts region properties using `regionprops`. The labels and
        region properties are stored in `self.labels_dict` and `self.regions_dict`
        respectively, with the category name as the key.

        For each category, the number of detected regions is printed to the console.
        """
        # iterate over each category in self.masks assign to the local variable cat, and mask --hint use the items() method
        ...

            # label the mask using the label function from skimage.measure
            # assign the output to the local variable lbl
            ...
            
            # extract the region properties using the regionprops function from skimage.measure
            # assign the output to the local variable regs
            ...
            
            # assign the labels and regions to the labels_dict and regions_dict attributes, with the category name as the key
            ...
            
            # print the number of regions detected for the category, the format is "[INFO] Detected <lenght(regs)> regions for <category>"
            ...

```

```python
grader.check("MaskBuilder-main-for-pig-wood-gray-blue")
```

```python
# Build the Visualizer class
...
    # Define the display_masks method, which takes in the masks and categories, add static to the method as a decorator
    ...
        """
        Display the masks for all categories.
        """
        # Create a figure and a set of subplots
        ...
        # Iterate over the categories and display the masks
        ...
            # Display the mask
            ...
            # Set the title of the mask
            ...
            # Turn off the axis
            ...
        # Show the plot
        ...
    # Define the display_overlays method, which takes in the labels_dict, gray_image, and categories. Add static to the method as a decorator
    ...
        """
        Display overlays for each category.
        """
        # Create a figure and a set of subplots 
        ...
        # Iterate over the categories and display the overlays
        ...
            # Create the overlay
            ...
            # Display the overlay
            ...
            # Set the title of the overlay
            ...
            # Turn off the axis
            ...
        # Show the plot
        ...
    # Define the display_bounding_boxes method, which takes in the image, regions_dict, and categories. Add static to the method as a decorator
    ...
        """
        Display bounding boxes on the original image.
        """
        # Create a color map for the categories as a dictionary as pig-> red, wood-> orange, gray-> gray, blue-> blue
        ...
        # Create a figure and a set of subplots
        ...
        # Display the image
        ...
        # Set the title of the image
        ...
        # Iterate over the categories and display the bounding boxes
        ...
            # Get the regions for the category
            ...
            # Get the color for the category
            ...
            # Iterate over the regions and display the bounding boxes
            ...
                # Check if the region area is greater than 50
                ...
                    # Get the bounding box for the region
                    ...
                    # Create a rectangle for the bounding box
                    ...
                    # Add the rectangle to the plot
                    ...
        # Turn off the axis
        ...
        # Show the plot
        ...
```

```python
grader.check("Visualizer-main")
```

```python
# image = io.imread("angry_birds_level.png")  # Replace with your local filename
image = structure_only  # Example image variable

# Initialize Mask Builder
mask_builder = MaskBuilder(image)

# Display Original Image
mask_builder.display_image(title="Original Image")

# Build and Clean Masks
mask_builder.build_and_clean_masks()

# Label and Extract Regions
mask_builder.label_and_extract_regions()

# Display Results
Visualizer.display_masks(mask_builder.masks, mask_builder.categories)
Visualizer.display_overlays(mask_builder.labels_dict, mask_builder.gray_image, mask_builder.categories)
Visualizer.display_bounding_boxes(image, mask_builder.regions_dict, mask_builder.categories)
```

```python
grader.check("writing-main")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week9-lab", "1_angry_birds_image")
```

Tests for Color-mask-and-merge (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(13.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(105.0)
>>> log_variable('total-points', f'week9-lab, 1_angry_birds_image', 105.0)
>>> question_id = 'Color-mask-and-merge-1'
>>> max_score = 2.0
>>> score = 0
>>> processor = ImageProcessor(image)
>>> color_mask = ColorMask(processor.hsv_image, (0.1, 0.5), (0.2, 0.8), (0.3, 0.9))
>>> color_mask_initialized = isinstance(color_mask, ColorMask)
>>> assert color_mask_initialized, 'ColorMask class did not initialize correctly.'
>>> if color_mask_initialized:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(color_mask_initialized))
```
failure_message: Failed: ColorMask class initialization failed.
success_message: Success: ColorMask class initialized successfully.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Color-mask-and-merge-2'
>>> max_score = 2.0
>>> score = 0
>>> processor = ImageProcessor(image)
>>> color_mask = ColorMask(processor.hsv_image, (0.1, 0.5), (0.2, 0.8), (0.3, 0.9))
>>> mask = color_mask.build_mask()
>>> mask_built = isinstance(mask, np.ndarray)
>>> assert mask_built, 'Mask was not built correctly.'
>>> if mask_built:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(mask_built))
```
failure_message: Failed: Mask building did not execute properly.
success_message: Success: Mask is built correctly.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Color-mask-and-merge-3'
>>> max_score = 3.0
>>> score = 0
>>> processor = ImageProcessor(image)
>>> color_mask = ColorMask(processor.hsv_image, (0.1, 0.5), (0.2, 0.8), (0.3, 0.9))
>>> mask = color_mask.build_mask()
>>> mask_boolean = mask.dtype == bool
>>> assert mask_boolean, 'Mask values are not boolean.'
>>> if mask_boolean:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(mask_boolean))
```
failure_message: Failed: Mask values are not boolean.
success_message: Success: Mask values are boolean.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Color-mask-and-merge-4'
>>> max_score = 3.0
>>> score = 0
>>> processor = ImageProcessor(image)
>>> color_mask = ColorMask(processor.hsv_image, (0.1, 0.5), (0.2, 0.8), (0.3, 0.9))
>>> mask = color_mask.build_mask()
>>> cleaned_mask = color_mask.clean_mask()
>>> mask_cleaned = isinstance(cleaned_mask, np.ndarray)
>>> assert mask_cleaned, 'Mask cleaning failed.'
>>> if mask_cleaned:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(mask_cleaned))
```
failure_message: Failed: Mask cleaning did not execute properly.
success_message: Success: Mask cleaned successfully.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Color-mask-and-merge-5'
>>> max_score = 3.0
>>> score = 0
>>> processor = ImageProcessor(image)
>>> color_mask = ColorMask(processor.hsv_image, (0.1, 0.5), (0.2, 0.8), (0.3, 0.9))
>>> mask = color_mask.build_mask()
>>> cleaned_mask = color_mask.clean_mask()
>>> cleaned_mask_boolean = cleaned_mask.dtype == bool
>>> assert cleaned_mask_boolean, 'Cleaned mask values are not boolean.'
>>> if cleaned_mask_boolean:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(cleaned_mask_boolean))
```
failure_message: Failed: Cleaned mask values are not boolean.
success_message: Success: Cleaned mask values remain boolean.


Tests for MaskBuilder-main-for-pig-wood-gray-blue (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(10.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(105.0)
>>> log_variable('total-points', f'week9-lab, 1_angry_birds_image', 105.0)
>>> question_id = 'MaskBuilder-main-for-pig-wood-gray-blue-1'
>>> max_score = 1.0
>>> score = 0
>>> mask_builder_initialized = isinstance(MaskBuilder(image), MaskBuilder)
>>> assert mask_builder_initialized, 'MaskBuilder did not initialize correctly.'
>>> if mask_builder_initialized:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(mask_builder_initialized))
```
failure_message: Failed: MaskBuilder initialization encountered an issue.
success_message: Success: MaskBuilder initializes with no errors.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'MaskBuilder-main-for-pig-wood-gray-blue-2'
>>> max_score = 1.0
>>> score = 0
>>> mask_builder = MaskBuilder(image)
>>> mask_builder.build_and_clean_masks()
>>> masks_built_successfully = isinstance(mask_builder.masks, dict) and len(mask_builder.masks) > 0
>>> assert masks_built_successfully, 'build_and_clean_masks did not execute properly.'
>>> if masks_built_successfully:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(masks_built_successfully))
```
failure_message: Failed: build_and_clean_masks encountered an issue.
success_message: Success: build_and_clean_masks method executes without errors.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'MaskBuilder-main-for-pig-wood-gray-blue-3'
>>> max_score = 4.0
>>> score = 0
>>> mask_builder = MaskBuilder(image)
>>> mask_builder.build_and_clean_masks()
>>> expected_categories = set(mask_builder.categories)
>>> actual_categories = set(mask_builder.masks.keys())
>>> mask_keys_correct = expected_categories == actual_categories
>>> assert mask_keys_correct, f'Expected mask keys: {expected_categories}, but got: {actual_categories}'
>>> if mask_keys_correct:
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(mask_keys_correct))
```
failure_message: Failed: Some masks do not have expected keys.
success_message: Success: Mask dictionary contains all expected keys.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from skimage.measure import regionprops
>>> question_id = 'MaskBuilder-main-for-pig-wood-gray-blue-4'
>>> max_score = 4.0
>>> score = 0
>>> mask_builder = MaskBuilder(image)
>>> mask_builder.build_and_clean_masks()
>>> valid_region_properties = all((isinstance(regs, list) and all((hasattr(r, 'area') and hasattr(r, 'bbox') for r in regs)) for regs in mask_builder.regions_dict.values()))
>>> assert valid_region_properties, 'Extracted regions do not contain valid region properties.'
>>> if valid_region_properties:
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(valid_region_properties))
```
failure_message: Failed: Some extracted regions do not contain valid properties.
success_message: Success: Regions extracted contain valid region properties.


Tests for PigMask-main (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(8.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(105.0)
>>> log_variable('total-points', f'week9-lab, 1_angry_birds_image', 105.0)
>>> question_id = 'PigMask-main-1'
>>> max_score = 2.0
>>> score = 0
>>> processor = ImageProcessor(image)
>>> pig_mask = PigMask(processor.hsv_image)
>>> pig_mask_initialized = isinstance(pig_mask, PigMask)
>>> assert pig_mask_initialized, 'PigMask class did not initialize correctly.'
>>> if pig_mask_initialized:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(pig_mask_initialized))
```
failure_message: Failed: PigMask class initialization failed.
success_message: Success: PigMask class initialized successfully.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'PigMask-main-2'
>>> max_score = 2.0
>>> score = 0
>>> processor = ImageProcessor(image)
>>> wood_mask = WoodMask(processor.hsv_image)
>>> wood_mask_initialized = isinstance(wood_mask, WoodMask)
>>> assert wood_mask_initialized, 'WoodMask class did not initialize correctly.'
>>> if wood_mask_initialized:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(wood_mask_initialized))
```
failure_message: Failed: WoodMask class initialization failed.
success_message: Success: WoodMask class initialized successfully.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'PigMask-main-3'
>>> max_score = 2.0
>>> score = 0
>>> processor = ImageProcessor(image)
>>> gray_mask = GrayMask(processor.hsv_image)
>>> gray_mask_initialized = isinstance(gray_mask, GrayMask)
>>> assert gray_mask_initialized, 'GrayMask class did not initialize correctly.'
>>> if gray_mask_initialized:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(gray_mask_initialized))
```
failure_message: Failed: GrayMask class initialization failed.
success_message: Success: GrayMask class initialized successfully.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'PigMask-main-4'
>>> max_score = 2.0
>>> score = 0
>>> processor = ImageProcessor(image)
>>> blue_mask = BlueMask(processor.hsv_image)
>>> blue_mask_initialized = isinstance(blue_mask, BlueMask)
>>> assert blue_mask_initialized, 'BlueMask class did not initialize correctly.'
>>> if blue_mask_initialized:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(blue_mask_initialized))
```
failure_message: Failed: BlueMask class initialization failed.
success_message: Success: BlueMask class initialized successfully.


Tests for Visualizer-main (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(8.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(105.0)
>>> log_variable('total-points', f'week9-lab, 1_angry_birds_image', 105.0)
>>> question_id = 'Visualizer-main-1'
>>> max_score = 1.0
>>> score = 0
>>> visualizer_initialized = isinstance(Visualizer(), Visualizer)
>>> assert visualizer_initialized, 'Visualizer class did not initialize correctly.'
>>> if visualizer_initialized:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(visualizer_initialized))
```
failure_message: Failed: Visualizer did not initialize correctly.
success_message: Success: Visualizer initializes correctly.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Visualizer-main-2'
>>> max_score = 1.0
>>> score = 0
>>> display_masks_success = callable(Visualizer.display_masks)
>>> assert display_masks_success, 'display_masks method is not callable.'
>>> if display_masks_success:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(display_masks_success))
```
failure_message: Failed: display_masks encountered an issue.
success_message: Success: display_masks method executes properly.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Visualizer-main-3'
>>> max_score = 1.0
>>> score = 0
>>> display_overlays_success = callable(Visualizer.display_overlays)
>>> assert display_overlays_success, 'display_overlays method is not callable.'
>>> if display_overlays_success:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(display_overlays_success))
```
failure_message: Failed: display_overlays encountered an issue.
success_message: Success: display_overlays method executes properly.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Visualizer-main-4'
>>> max_score = 1.0
>>> score = 0
>>> display_bounding_boxes_success = callable(Visualizer.display_bounding_boxes)
>>> assert display_bounding_boxes_success, 'display_bounding_boxes method is not callable.'
>>> if display_bounding_boxes_success:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(display_bounding_boxes_success))
```
failure_message: Failed: display_bounding_boxes encountered an issue.
success_message: Success: display_bounding_boxes method executes properly.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Visualizer-main-5'
>>> max_score = 4.0
>>> score = 0
>>> visualizer_methods_correct = all((hasattr(Visualizer, method) for method in ['display_masks', 'display_overlays', 'display_bounding_boxes']))
>>> assert visualizer_methods_correct, 'Some static methods in Visualizer are missing or incorrect.'
>>> if visualizer_methods_correct:
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(visualizer_methods_correct))
```
failure_message: Failed: Some static methods in Visualizer are missing or incorrect.
success_message: Success: All static methods in Visualizer are correctly implemented.


Tests for edge-detection (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> import matplotlib.pyplot as plt
>>> from skimage import io, color, filters, feature, morphology
>>> from skimage.measure import label, regionprops
>>> from skimage.segmentation import clear_border
>>> max_question_points = str(10.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(105.0)
>>> log_variable('total-points', f'week9-lab, 1_angry_birds_image', 105.0)
>>> question_id = 'edge-detection-1'
>>> max_score = 1.0
>>> score = 0
>>> edges_type = isinstance(edges, np.ndarray)
>>> assert edges_type
>>> if edges_type:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(edges_type))
```
failure_message: Failed: Edge detection did not execute properly.
success_message: Success: Canny edge detection applied successfully.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> import matplotlib.pyplot as plt
>>> from skimage import io, color, filters, feature, morphology
>>> from skimage.measure import label, regionprops
>>> from skimage.segmentation import clear_border
>>> question_id = 'edge-detection-2'
>>> max_score = 1.0
>>> score = 0
>>> edges_dtype = edges.dtype == bool
>>> assert edges_dtype, 'Edge detection result should be a boolean array.'
>>> if edges_dtype:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(edges_dtype))
```
failure_message: Failed: Edge matrix does not contain only boolean values.
success_message: Success: Edge matrix contains boolean values.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> import matplotlib.pyplot as plt
>>> from skimage import io, color, filters, feature, morphology
>>> from skimage.measure import label, regionprops
>>> from skimage.segmentation import clear_border
>>> question_id = 'edge-detection-3'
>>> max_score = 1.0
>>> score = 0
>>> edges_shape_match = edges.shape == gray_image.shape
>>> assert edges_shape_match, 'Edge detection output shape mismatch.'
>>> if edges_shape_match:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(edges_shape_match))
```
failure_message: Failed: Edge detection output shape does not match input shape.
success_message: Success: Edge detection output shape matches input shape.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> import matplotlib.pyplot as plt
>>> from skimage import io, color, filters, feature, morphology
>>> from skimage.measure import label, regionprops
>>> from skimage.segmentation import clear_border
>>> question_id = 'edge-detection-4'
>>> max_score = 1.0
>>> score = 0
>>> edges_detected = np.any(edges)
>>> assert edges_detected, 'No edges were detected in the image.'
>>> if edges_detected:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(edges_detected))
```
failure_message: Failed: No edges detected in the image.
success_message: Success: At least some edges detected.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> import matplotlib.pyplot as plt
>>> from skimage import io, color, filters, feature, morphology
>>> from skimage.measure import label, regionprops
>>> from skimage.segmentation import clear_border
>>> question_id = 'edge-detection-5'
>>> max_score = 1.0
>>> score = 0
>>> edges_detected = np.any(edges)
>>> assert edges_detected, 'No edges were detected in the image.'
>>> if edges_detected:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(edges_detected))
```
failure_message: Failed: No edges detected in the image.
success_message: Success: At least some edges detected.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> import matplotlib.pyplot as plt
>>> from skimage import io, color, filters, feature, morphology
>>> from skimage.measure import label, regionprops
>>> from skimage.segmentation import clear_border
>>> question_id = 'edge-detection-6'
>>> max_score = 5.0
>>> score = 0
>>> exec(base64.b64decode('ZWRnZXNfZGV0ZWN0ZWQgPSBucC5zdW0oZWRnZXMp').decode())
>>> assert edges_detected == 28996, 'Edges detected incorrectly.'
>>> if edges_detected == 28996:
...     score = 5.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(edges_detected))
```
failure_message: Failed: Edges detected incorrectly.
success_message: Success: Edges detected correctly.


Tests for image-processor (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(14.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(105.0)
>>> log_variable('total-points', f'week9-lab, 1_angry_birds_image', 105.0)
>>> question_id = 'image-processor-1'
>>> max_score = 2.0
>>> score = 0
>>> processor = ImageProcessor(image)
>>> processor_initialized = isinstance(processor, ImageProcessor)
>>> assert processor_initialized, 'ImageProcessor class did not initialize correctly.'
>>> if processor_initialized:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(processor_initialized))
```
failure_message: Failed: ImageProcessor class initialization failed.
success_message: Success: ImageProcessor class initialized successfully.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'image-processor-2'
>>> max_score = 2.0
>>> score = 0
>>> processor = ImageProcessor(image)
>>> image_attribute_type = isinstance(processor.image, np.ndarray)
>>> assert image_attribute_type, 'Image attribute is not a NumPy array.'
>>> if image_attribute_type:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(image_attribute_type))
```
failure_message: Failed: Image attributes are not assigned correctly.
success_message: Success: Image attributes are correctly assigned.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'image-processor-3'
>>> max_score = 2.0
>>> score = 0
>>> processor = ImageProcessor(image)
>>> gray_image_type = isinstance(processor.gray_image, np.ndarray)
>>> assert gray_image_type, 'Grayscale image computation failed.'
>>> if gray_image_type:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(gray_image_type))
```
failure_message: Failed: Grayscale image is not computed correctly.
success_message: Success: Grayscale image is computed correctly.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'image-processor-4'
>>> max_score = 2.0
>>> score = 0
>>> processor = ImageProcessor(image)
>>> hsv_image_type = isinstance(processor.hsv_image, np.ndarray)
>>> assert hsv_image_type, 'HSV image computation failed.'
>>> if hsv_image_type:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(hsv_image_type))
```
failure_message: Failed: HSV image is not computed correctly.
success_message: Success: HSV image is computed correctly.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'image-processor-5'
>>> max_score = 2.0
>>> score = 0
>>> required_methods = ['__init__', 'display_image']
>>> methods_exist = all((hasattr(ImageProcessor, method) for method in required_methods))
>>> assert methods_exist, 'One or more required methods are missing in the ImageProcessor class.'
>>> if methods_exist:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(methods_exist))
```
failure_message: Failed: Some required methods are missing in the ImageProcessor class.
success_message: Success: ImageProcessor class has the required methods.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'image-processor-6'
>>> max_score = 4.0
>>> score = 0
>>> processor = ImageProcessor(image)
>>> required_attributes = ['image', 'hsv_image', 'gray_image', 'masks', 'labels_dict', 'regions_dict']
>>> attributes_exist = all((hasattr(processor, attr) for attr in required_attributes))
>>> assert attributes_exist, 'One or more required attributes are missing in the ImageProcessor class.'
>>> if attributes_exist:
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(attributes_exist))
```
failure_message: Failed: Some required attributes are missing in the ImageProcessor class.
success_message: Success: ImageProcessor attributes are correctly defined.


Tests for import_packages (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> max_question_points = str(3.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(105.0)
>>> log_variable('total-points', f'week9-lab, 1_angry_birds_image', 105.0)
>>> question_id = 'import_packages-1'
>>> max_score = 1.0
>>> score = 0
>>> numpy_imported = 'np' in globals()
>>> assert numpy_imported, 'numpy is not imported with alias np'
>>> if numpy_imported:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(numpy_imported))
```
failure_message: Failed: numpy is not imported correctly.
success_message: Success: numpy is imported correctly.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import matplotlib.pyplot as plt
>>> question_id = 'import_packages-2'
>>> max_score = 1.0
>>> score = 0
>>> matplotlib_imported = 'plt' in globals()
>>> assert matplotlib_imported, 'matplotlib.pyplot is not imported with alias plt'
>>> if matplotlib_imported:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(matplotlib_imported))
```
failure_message: Failed: matplotlib.pyplot is not imported correctly.
success_message: Success: matplotlib.pyplot is imported correctly.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from skimage import io, color, filters, feature, morphology
>>> from skimage.measure import label, regionprops
>>> from skimage.segmentation import clear_border
>>> question_id = 'import_packages-3'
>>> max_score = 1.0
>>> score = 0
>>> skimage_modules = ['io', 'color', 'filters', 'feature', 'morphology']
>>> skimage_measure_modules = ['label', 'regionprops']
>>> skimage_segmentation_modules = ['clear_border']
>>> skimage_imported = all((mod in globals() for mod in skimage_modules + skimage_measure_modules + skimage_segmentation_modules))
>>> assert skimage_imported, 'Some skimage modules are missing from the import statement'
>>> if skimage_imported:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(skimage_imported))
```
failure_message: Failed: skimage modules are not imported correctly.
success_message: Success: skimage modules are imported correctly.


Tests for isolate-main-structure (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> import matplotlib.pyplot as plt
>>> from skimage import io, color, filters, feature, morphology
>>> from skimage.measure import label, regionprops
>>> from skimage.segmentation import clear_border
>>> max_question_points = str(12.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(105.0)
>>> log_variable('total-points', f'week9-lab, 1_angry_birds_image', 105.0)
>>> question_id = 'isolate-main-structure-1'
>>> max_score = 1.0
>>> score = 0
>>> structure_only_type = isinstance(structure_only, np.ndarray)
>>> assert structure_only_type, 'Image copy for structure isolation failed.'
>>> if structure_only_type:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(structure_only_type))
```
failure_message: Failed: Image was not copied correctly.
success_message: Success: Image copied correctly for structure isolation.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> import matplotlib.pyplot as plt
>>> from skimage import io, color, filters, feature, morphology
>>> from skimage.measure import label, regionprops
>>> from skimage.segmentation import clear_border
>>> question_id = 'isolate-main-structure-2'
>>> max_score = 1.0
>>> score = 0
>>> inverted_mask_type = isinstance(inverted_mask, np.ndarray)
>>> assert inverted_mask_type, 'Mask inversion failed.'
>>> if inverted_mask_type:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(inverted_mask_type))
```
failure_message: Failed: Mask inversion did not execute properly.
success_message: Success: Mask was inverted successfully.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> import matplotlib.pyplot as plt
>>> from skimage import io, color, filters, feature, morphology
>>> from skimage.measure import label, regionprops
>>> from skimage.segmentation import clear_border
>>> question_id = 'isolate-main-structure-3'
>>> max_score = 2.0
>>> score = 0
>>> background_modified = np.any(structure_only == 0)
>>> assert background_modified, 'Background setting failed.'
>>> if background_modified:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(background_modified))
```
failure_message: Failed: Background was not modified correctly.
success_message: Success: Background set to white correctly.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> import matplotlib.pyplot as plt
>>> from skimage import io, color, filters, feature, morphology
>>> from skimage.measure import label, regionprops
>>> from skimage.segmentation import clear_border
>>> question_id = 'isolate-main-structure-4'
>>> max_score = 1.0
>>> score = 0
>>> fig_created = isinstance(fig, plt.Figure)
>>> assert fig_created, 'Figure creation failed.'
>>> if fig_created:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(fig_created))
```
failure_message: Failed: Visualization figure creation failed.
success_message: Success: Visualization figure created successfully.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> import matplotlib.pyplot as plt
>>> from skimage import io, color, filters, feature, morphology
>>> from skimage.measure import label, regionprops
>>> from skimage.segmentation import clear_border
>>> question_id = 'isolate-main-structure-5'
>>> max_score = 3.0
>>> score = 0
>>> assert image.shape == (640, 960, 3), 'Image dimensions mismatch.'
>>> if image.shape == (640, 960, 3):
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(image))
```
failure_message: Failed: Image dimensions do not match original.
success_message: Success: Image dimensions match original image dimensions.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> import matplotlib.pyplot as plt
>>> from skimage import io, color, filters, feature, morphology
>>> from skimage.measure import label, regionprops
>>> from skimage.segmentation import clear_border
>>> question_id = 'isolate-main-structure-6'
>>> max_score = 4.0
>>> score = 0
>>> structure_shape_match = structure_only.shape == image.shape
>>> assert structure_shape_match, 'Structure-only image dimensions mismatch.'
>>> if structure_shape_match:
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(structure_shape_match))
```
failure_message: Failed: Structure-only image dimensions do not match original.
success_message: Success: Structure-only image dimensions match original image dimensions.


Tests for load-and-visualize-image (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> import matplotlib.pyplot as plt
>>> from skimage import io, color, filters, feature, morphology
>>> from skimage.measure import label, regionprops
>>> from skimage.segmentation import clear_border
>>> max_question_points = str(9.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(105.0)
>>> log_variable('total-points', f'week9-lab, 1_angry_birds_image', 105.0)
>>> question_id = 'load-and-visualize-image-1'
>>> max_score = 1.0
>>> score = 0
>>> assert image_path == 'angry-bird.webp', 'Image path is incorrect.'
>>> if image_path == 'angry-bird.webp':
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(image_path))
```
failure_message: Failed: Image path is incorrect.
success_message: Success: Image path assigned correctly.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> import matplotlib.pyplot as plt
>>> from skimage import io, color, filters, feature, morphology
>>> from skimage.measure import label, regionprops
>>> from skimage.segmentation import clear_border
>>> question_id = 'load-and-visualize-image-2'
>>> max_score = 1.0
>>> score = 0
>>> image_loaded = isinstance(image, np.ndarray)
>>> assert image_loaded, 'Image failed to load.'
>>> if image_loaded:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(image_loaded))
```
failure_message: Failed: Image not loaded correctly.
success_message: Success: Image loaded successfully.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> import matplotlib.pyplot as plt
>>> from skimage import io, color, filters, feature, morphology
>>> from skimage.measure import label, regionprops
>>> from skimage.segmentation import clear_border
>>> question_id = 'load-and-visualize-image-3'
>>> max_score = 1.0
>>> score = 0
>>> gray_image_type = isinstance(gray_image, np.ndarray)
>>> assert gray_image_type, 'Grayscale conversion failed.'
>>> if gray_image_type:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(gray_image_type))
```
failure_message: Failed: Image not converted to grayscale correctly.
success_message: Success: Image converted to grayscale.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> import matplotlib.pyplot as plt
>>> from skimage import io, color, filters, feature, morphology
>>> from skimage.measure import label, regionprops
>>> from skimage.segmentation import clear_border
>>> question_id = 'load-and-visualize-image-4'
>>> max_score = 1.0
>>> score = 0
>>> fig_created = isinstance(fig, plt.Figure)
>>> assert fig_created, 'Figure creation failed.'
>>> if fig_created:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(fig_created))
```
failure_message: Failed: Image display did not execute properly.
success_message: Success: Image displayed successfully.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> import matplotlib.pyplot as plt
>>> from skimage import io, color, filters, feature, morphology
>>> from skimage.measure import label, regionprops
>>> from skimage.segmentation import clear_border
>>> question_id = 'load-and-visualize-image-5'
>>> max_score = 5.0
>>> score = 0
>>> exec(base64.b64decode('dmFsdWUgPSBucC5zdW0oaW1hZ2Up').decode())
>>> assert value == 329469408, 'Figure creation failed.'
>>> if value == 329469408:
...     score = 5.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(image))
```
failure_message: Failed: Image display did not execute properly.
success_message: Success: Image displayed successfully.


Tests for refinement-dilation-and-hole-filling (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> import matplotlib.pyplot as plt
>>> from skimage import io, color, filters, feature, morphology
>>> from skimage.measure import label, regionprops
>>> from skimage.segmentation import clear_border
>>> max_question_points = str(13.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(105.0)
>>> log_variable('total-points', f'week9-lab, 1_angry_birds_image', 105.0)
>>> question_id = 'refinement-dilation-and-hole-filling-1'
>>> max_score = 1.0
>>> score = 0
>>> dilated_edges_type = isinstance(dilated_edges, np.ndarray)
>>> assert dilated_edges_type, 'Edge dilation failed.'
>>> if dilated_edges_type:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(dilated_edges_type))
```
failure_message: Failed: Edge dilation did not execute properly.
success_message: Success: Edges dilated correctly.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> import matplotlib.pyplot as plt
>>> from skimage import io, color, filters, feature, morphology
>>> from skimage.measure import label, regionprops
>>> from skimage.segmentation import clear_border
>>> question_id = 'refinement-dilation-and-hole-filling-2'
>>> max_score = 1.0
>>> score = 0
>>> filled_mask_type = isinstance(filled_mask, np.ndarray)
>>> assert filled_mask_type, 'Hole filling failed.'
>>> if filled_mask_type:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(filled_mask_type))
```
failure_message: Failed: Hole filling did not execute properly.
success_message: Success: Holes filled correctly.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> import matplotlib.pyplot as plt
>>> from skimage import io, color, filters, feature, morphology
>>> from skimage.measure import label, regionprops
>>> from skimage.segmentation import clear_border
>>> question_id = 'refinement-dilation-and-hole-filling-3'
>>> max_score = 1.0
>>> score = 0
>>> filled_mask_cleaned = np.any(filled_mask)
>>> assert filled_mask_cleaned, 'Filled mask is empty after small object removal.'
>>> if filled_mask_cleaned:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(filled_mask_cleaned))
```
failure_message: Failed: Small objects were not removed properly.
success_message: Success: Small objects removed correctly.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> import matplotlib.pyplot as plt
>>> from skimage import io, color, filters, feature, morphology
>>> from skimage.measure import label, regionprops
>>> from skimage.segmentation import clear_border
>>> question_id = 'refinement-dilation-and-hole-filling-4'
>>> max_score = 1.0
>>> score = 0
>>> fig_created = isinstance(fig, plt.Figure)
>>> assert fig_created, 'Figure creation failed.'
>>> if fig_created:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(fig_created))
```
failure_message: Failed: Visualization figure creation failed.
success_message: Success: Visualization figure created successfully.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> import matplotlib.pyplot as plt
>>> from skimage import io, color, filters, feature, morphology
>>> from skimage.measure import label, regionprops
>>> from skimage.segmentation import clear_border
>>> question_id = 'refinement-dilation-and-hole-filling-5'
>>> max_score = 3.0
>>> score = 0
>>> assert filled_mask.shape == (640, 960), 'Filled mask dimensions do not match original image dimensions.'
>>> if filled_mask.shape == (640, 960):
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(filled_mask))
```
failure_message: Failed: Filled mask dimensions do not match input image dimensions.
success_message: Success: Filled mask dimensions match input image dimensions.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> import matplotlib.pyplot as plt
>>> from skimage import io, color, filters, feature, morphology
>>> from skimage.measure import label, regionprops
>>> from skimage.segmentation import clear_border
>>> question_id = 'refinement-dilation-and-hole-filling-6'
>>> max_score = 3.0
>>> score = 0
>>> assert edges.shape == (640, 960), 'Edge dimensions do not match original image dimensions.'
>>> if edges.shape == (640, 960):
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(edges))
```
failure_message: Failed: Edge dimensions do not match input image dimensions.
success_message: Success: Edge dimensions match input image dimensions.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> import matplotlib.pyplot as plt
>>> from skimage import io, color, filters, feature, morphology
>>> from skimage.measure import label, regionprops
>>> from skimage.segmentation import clear_border
>>> question_id = 'refinement-dilation-and-hole-filling-7'
>>> max_score = 3.0
>>> score = 0
>>> mask_shape_match = filled_mask.shape == edges.shape
>>> assert mask_shape_match, 'Filled mask dimensions do not match original image dimensions.'
>>> if mask_shape_match:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(mask_shape_match))
```
failure_message: Failed: Mask dimensions do not match input image dimensions.
success_message: Success: Mask dimensions match input image dimensions.


Tests for writing-main (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(5.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(105.0)
>>> log_variable('total-points', f'week9-lab, 1_angry_birds_image', 105.0)
>>> question_id = 'writing-main-1'
>>> max_score = 1.0
>>> score = 0
>>> pipeline_executed = True
>>> assert pipeline_executed, 'MaskBuilder pipeline did not execute correctly.'
>>> if pipeline_executed:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(pipeline_executed))
```
failure_message: Failed: MaskBuilder pipeline encountered an issue.
success_message: Success: MaskBuilder pipeline runs correctly.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'writing-main-2'
>>> max_score = 1.0
>>> score = 0
>>> masks_built = isinstance(mask_builder.masks, dict) and len(mask_builder.masks) > 0
>>> assert masks_built, 'Masks were not created properly.'
>>> if masks_built:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(masks_built))
```
failure_message: Failed: Masks were not created properly.
success_message: Success: Masks are built and stored properly.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'writing-main-3'
>>> max_score = 1.0
>>> score = 0
>>> labels_extracted = isinstance(mask_builder.labels_dict, dict) and len(mask_builder.labels_dict) > 0
>>> assert labels_extracted, 'Labels and regions were not extracted properly.'
>>> if labels_extracted:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(labels_extracted))
```
failure_message: Failed: Labels and regions were not extracted properly.
success_message: Success: Labels and regions extracted properly.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'writing-main-4'
>>> max_score = 1.0
>>> score = 0
>>> visualizer_functions = all((callable(getattr(Visualizer, method)) for method in ['display_masks', 'display_overlays', 'display_bounding_boxes']))
>>> assert visualizer_functions, 'One or more Visualizer functions did not execute correctly.'
>>> if visualizer_functions:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(visualizer_functions))
```
failure_message: Failed: One or more Visualizer functions did not execute correctly.
success_message: Success: Visualizer functions execute correctly.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'writing-main-5'
>>> max_score = 1.0
>>> score = 0
>>> bounding_boxes_displayed = True
>>> assert bounding_boxes_displayed, 'Bounding boxes were not displayed correctly.'
>>> if bounding_boxes_displayed:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_angry_birds_image', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(bounding_boxes_displayed))
```
failure_message: Failed: Bounding boxes were not displayed correctly.
success_message: Success: Bounding boxes displayed correctly.
