```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("4_scikit_image_q", "week_9", "readings", assignment_points = 45.0, assignment_tag = 'week9-readings')

# Initialize Otter
import otter
grader = otter.Notebook("4_scikit_image_q.ipynb")
```

# ❓🧩 Survive in the World into Pixels of Interest

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._4_scikit_image_q import Question1
Question1().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._4_scikit_image_q import Question2
Question2().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._4_scikit_image_q import Question3
Question3().show()

```

## Question 1 (Points: 14.0): Tipu's Tiger Mask Analysis

### Instructions:
Tipu's Tiger is an 18th-century mechanical automaton and pipe organ created during the reign of Tipu Sultan, the ruler of Mysore, symbolizing his resistance against British colonial forces. It depicts a tiger mauling a European soldier, with a crank-operated mechanism that produces tiger growls, the soldier’s cries, and music from a hidden organ that can be played via a keyboard. Captured after Tipu’s death in 1799 during the fall of Seringapatam, it was taken to England as a war trophy and is now housed in the Victoria and Albert Museum in London. This artifact remains a powerful symbol of Tipu’s defiance, his kingdom's technological ingenuity, and cultural sophistication. 

<iframe width="560" height="315" src="https://www.youtube.com/embed/XNJRjV8K1yg?si=rv_vdBMtzfO7RQEn" title="YouTube video player" frameborder="0" allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share" referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>

Suppose you are tasked with converting the background of an image of this artifact to match Drexel branding. For this you want to isolate the tiger from the background using masking techniques in `scikit-image`, and replace the black background with the Drexel blue color [0, 47, 108] Complete the following:

Your job is to implement the process_tipu_tiger() function that takes the image as input and return the original image, the modified image, mask, and the modified image with the Drexel blue background. We have provided the scaffold for the function.

1. Load the image "tipu-tiger.jpg" from the folder, and store it in a variable named `image`. You can use the `ski.io.imread` function to load the image. 
2. Create a mask: Isolate pixels with intensity in a specific range (example: RGB channel sum > 150). Store the mask in a variable named `mask`. Hint: you should use the `numpy.sum` function to sum the RGB channels, making sure to do this along the 3rd axis. You can use the logical operator `>` to create the mask. 
3. Create a disk of radius 28 using the `morphology.disk` function from the `skimage` module. Store the disk in a variable named `selem`. Note: selem stands for "structuring element" used in morphological operations of images. We have already imported the necessary modules for you. This should provide a good hint.
4. Apply a dilation operation to the mask, and then an erosion operation to the dilated mask using the `morphology.dilation` and `morphology.erosion` functions from the `skimage` module. Store the dilated mask in a variable named `dilated_mask` and the eroded mask in a variable named `eroded_mask`. The process of dilation followed by erosion is known as closing in image processing.
5. Remove small holes using the `remove_small_holes` function from the `skimage.morphology` module. Store the result in a variable named `final_mask`. Make sure to set the `area_threshold` parameter to 5000.
6. Replace all pixels not in the masked region with the color black. This can be done by indexing the image with the inverse of the mask. Hint: you can use the ~ operator to invert the mask. A black image has RGB values of [0, 0, 0]. Store the modified image in a variable named `modified_image`.
7. Replace the black background with the Drexel blue color `[0, 47, 108]`. Store the modified image with the Drexel blue background in a variable named `modified_image_drexel`. This should be done similarly to step 6, but using the Drexel blue color instead of black.

When you run the code, you will see plots generated using the provided code.



```python
import skimage as ski
import numpy as np
from skimage.morphology import dilation, erosion, disk, remove_small_holes
import matplotlib.pyplot as plt


def process_tipu_tiger(image_path="tipu-tiger.jpg"):
    # Load the image

    # 1. Load the image "tipu-tiger.jpg" from the folder, and store it in a variable named `image`.
    # You can use the `ski.io.imread` function to load the image.
    ...

    # 2. Create a mask: Isolate pixels with intensity in a specific range (example: RGB channel sum > 150)
    ...

    # 3. Create the disk-shaped structuring element
    ...

    # 4. Apply the dilation and erosion operations
    ...

    # 5. Fill holes in the mask
    ...

    # makes a copy of the image
    # Do not modify this line
    modified_image = image.copy()

    # 6. Replace the background with black color
    ...

    # makes a copy of the resulting image
    modified_image_drexel = modified_image.copy()

    # 7. Replaces the background with the Drexel blue color
    ...

    return (
        image,
        modified_image,
        mask,
        final_mask,
        modified_image_drexel,
        dilated_mask,
        eroded_mask,
    )


# Example usage
(
    image,
    modified_image,
    mask,
    final_mask,
    modified_image_drexel,
    dilated_mask,
    eroded_mask,
) = process_tipu_tiger()


# Plot results
plt.figure(figsize=(12, 6))

plt.subplot(1, 5, 1)
plt.imshow(image)
plt.title("Image")
plt.axis("off")

plt.subplot(1, 5, 2)
plt.imshow(modified_image)
plt.title("Masked Image")
plt.axis("off")

plt.subplot(1, 5, 3)
plt.imshow(mask, cmap="gray")
plt.title("Mask")
plt.axis("off")

plt.subplot(1, 5, 4)
plt.imshow(final_mask, cmap="gray")
plt.title("Final Mask")
plt.axis("off")

plt.subplot(1, 5, 5)
plt.imshow(modified_image_drexel)
plt.title("Recolored Image")
plt.axis("off")

plt.tight_layout()
plt.show()
```

```python
grader.check("tipus-tiger-mask-calculation")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week9-readings", "4_scikit_image_q")
```

Tests for tipus-tiger-mask-calculation (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import re
>>> import inspect
>>> max_question_points = str(14.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(14.0)
>>> log_variable('total-points', f'week9-readings, 4_scikit_image_q', 14.0)
>>> question_id = 'tipus-tiger-mask-calculation-1'
>>> max_score = 1.0
>>> score = 0
>>> def get_process_tipu_tiger_code(process_tipu_tiger):
...     """
...     Returns the entire source code of process_tipu_tiger as a string.
...     """
...     return inspect.getsource(process_tipu_tiger)
>>> code = get_process_tipu_tiger_code(process_tipu_tiger)
>>> load_pattern = '.*imread\\(.*\\).*'
>>> load_pattern_found = bool(re.search(load_pattern, code))
>>> assert load_pattern_found, 'You must call `ski.io.imread` to load the image.'
>>> if load_pattern_found:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_scikit_image_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(load_pattern_found))
```
failure_message: Failed: Must load the image from 'tipu-tiger.jpg'.
success_message: Success: Used ski.io.imread with the correct path.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import re
>>> import inspect
>>> question_id = 'tipus-tiger-mask-calculation-2'
>>> max_score = 1.0
>>> score = 0
>>> def get_process_tipu_tiger_code(process_tipu_tiger):
...     """
...     Returns the entire source code of process_tipu_tiger as a string.
...     """
...     return inspect.getsource(process_tipu_tiger)
>>> code = get_process_tipu_tiger_code(process_tipu_tiger)
>>> mask_pattern = 'mask\\s*=\\s*np\\.sum\\s*\\(\\s*image\\s*,\\s*axis\\s*=\\s*2\\s*\\)\\s*>\\s*150'
>>> mask_pattern_found = bool(re.search(mask_pattern, code))
>>> assert mask_pattern_found, 'You must define `mask = np.sum(image, axis=2) > 150`.'
>>> if mask_pattern_found:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_scikit_image_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(mask_pattern_found))
```
failure_message: Failed: Must define mask = np.sum(image, axis=2) > 150.
success_message: Success: The mask uses np.sum(image, axis=2) > 150.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import re
>>> import inspect
>>> question_id = 'tipus-tiger-mask-calculation-3'
>>> max_score = 1.0
>>> score = 0
>>> def get_process_tipu_tiger_code(process_tipu_tiger):
...     """
...     Returns the entire source code of process_tipu_tiger as a string.
...     """
...     return inspect.getsource(process_tipu_tiger)
>>> code = get_process_tipu_tiger_code(process_tipu_tiger)
>>> disk_pattern = 'selem\\s*=\\s*disk\\s*\\(\\s*28\\s*\\)'
>>> disk_pattern_found = bool(re.search(disk_pattern, code))
>>> assert disk_pattern_found, 'You must define `selem = disk(28)`.'
>>> if disk_pattern_found:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_scikit_image_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(disk_pattern_found))
```
failure_message: Failed: Must define selem = disk(28).
success_message: Success: The code uses selem = disk(28).
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import re
>>> import inspect
>>> question_id = 'tipus-tiger-mask-calculation-4'
>>> max_score = 1.0
>>> score = 0
>>> def get_process_tipu_tiger_code(process_tipu_tiger):
...     """
...     Returns the entire source code of process_tipu_tiger as a string.
...     """
...     return inspect.getsource(process_tipu_tiger)
>>> code = get_process_tipu_tiger_code(process_tipu_tiger)
>>> dilation_pattern = 'dilated_mask\\s*=\\s*dilation\\s*\\(\\s*mask\\s*,\\s*selem\\s*\\)'
>>> erosion_pattern = 'eroded_mask\\s*=\\s*erosion\\s*\\(\\s*dilated_mask\\s*,\\s*selem\\s*\\)'
>>> dilation_pattern_found = bool(re.search(dilation_pattern, code))
>>> erosion_pattern_found = bool(re.search(erosion_pattern, code))
>>> assert dilation_pattern_found and erosion_pattern_found, 'You must define `dilated_mask = dilation(mask, selem)` then `eroded_mask = erosion(dilated_mask, selem)`.'
>>> if dilation_pattern_found and erosion_pattern_found:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_scikit_image_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(dilation_pattern_found))
>>> responses = update_responses(question_id, str(erosion_pattern_found))
```
failure_message: Failed: Must do dilated_mask = dilation(mask, selem) then eroded_mask = erosion(dilated_mask, selem).
success_message: Success: Dilation and erosion are applied in the correct order.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import re
>>> import inspect
>>> question_id = 'tipus-tiger-mask-calculation-5'
>>> max_score = 1.0
>>> score = 0
>>> def get_process_tipu_tiger_code(process_tipu_tiger):
...     """
...     Returns the entire source code of process_tipu_tiger as a string.
...     """
...     return inspect.getsource(process_tipu_tiger)
>>> code = get_process_tipu_tiger_code(process_tipu_tiger)
>>> remove_holes_pattern = 'final_mask\\s*=\\s*remove_small_holes\\s*\\(\\s*eroded_mask\\s*,\\s*area_threshold\\s*=\\s*\\d+\\s*\\)'
>>> remove_holes_pattern_found = bool(re.search(remove_holes_pattern, code))
>>> assert remove_holes_pattern_found, 'You must define `final_mask = remove_small_holes(eroded_mask, area_threshold=5000)` (or 50000).'
>>> if remove_holes_pattern_found:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_scikit_image_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(remove_holes_pattern_found))
```
failure_message: Failed: Must define final_mask = remove_small_holes(eroded_mask, area_threshold=...).
success_message: Success: remove_small_holes is called on eroded_mask with area_threshold.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import re
>>> import inspect
>>> question_id = 'tipus-tiger-mask-calculation-6'
>>> max_score = 1.0
>>> score = 0
>>> def get_process_tipu_tiger_code(process_tipu_tiger):
...     """
...     Returns the entire source code of process_tipu_tiger as a string.
...     """
...     return inspect.getsource(process_tipu_tiger)
>>> code = get_process_tipu_tiger_code(process_tipu_tiger)
>>> black_pattern = 'modified_image\\s*\\[\\s*~final_mask\\s*\\]\\s*=\\s*(?:0|\\[\\s*0\\s*,\\s*0\\s*,\\s*0\\s*\\])'
>>> black_pattern_found = bool(re.search(black_pattern, code))
>>> assert black_pattern_found, 'You must define either:\n  modified_image[~final_mask] = [0, 0, 0]\nor\n  modified_image[~final_mask] = 0\nto set the background to black.'
>>> if (black_pattern_found, 'You must define either:\n  modified_image[~final_mask] = [0, 0, 0]\nor\n  modified_image[~final_mask] = 0\nto set the background to black.'):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_scikit_image_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(black_pattern_found))
```
failure_message: Failed: Must replace the background with black using modified_image[~final_mask] = [0,0,0] or 0.
success_message: Success: The code sets modified_image[~final_mask] = [0, 0, 0] or 0.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import re
>>> import inspect
>>> question_id = 'tipus-tiger-mask-calculation-7'
>>> max_score = 2.0
>>> score = 0
>>> image_path = './tipu-tiger.jpg'
>>> _, _, mask, _, _, _, _ = process_tipu_tiger(image_path)
>>> sum_mask = np.sum(mask)
>>> assert sum_mask == 159187, 'The mask does not appear to be computed correctly.'
>>> if sum_mask == 159187:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_scikit_image_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(sum_mask))
```
failure_message: Mask does not appear to be computed correctly.
success_message: Success: Mask appears to be computed correctly.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import re
>>> import inspect
>>> question_id = 'tipus-tiger-mask-calculation-8'
>>> max_score = 6.0
>>> score = 0
>>> image_path = './tipu-tiger.jpg'
>>> _, _, _, _, modified_image_drexel_sum, _, _ = process_tipu_tiger()
>>> modified_image_drexel_sum = np.sum(modified_image_drexel.reshape(-1, 3), axis=0)
>>> condition1 = all(modified_image_drexel_sum == [22885055, 30374692, 41060560])
>>> assert condition1, 'The mask does not appear to be computed correctly.'
>>> if condition1:
...     score = 6.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_scikit_image_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(modified_image_drexel_sum))
```
failure_message: The final does not appear to be computed correctly.
success_message: Success: final image appears to be computed correctly.
