```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment(
    "10_scikit_learn_q",
    "week_9",
    "readings",
    assignment_points=43.0,
    assignment_tag="week9-readings",
)

# Initialize Otter
import otter

grader = otter.Notebook("10_scikit_learn_q.ipynb")
```

# â“ Machine Learning with scikit-learn

Note: we expect you use the sklearn documentation to answer these questions. Searching and reading documentation is a crucial skill for coding and machine learning. College is not about the material you learn but your ability to learn how to learn.

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._10_scikit_learn_q import Question1

Question1().show()
```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._10_scikit_learn_q import Question1

Question1().show()
```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._10_scikit_learn_q import Question2

Question2().show()
```

# Hello World of Machine Learning

The MNIST dataset is a collection of handwritten digits. Each image is 28x28 pixels, and each pixel is a grayscale value between 0 and 255. The dataset is split into a training set and a test set, with 60,000 images in the training set and 10,000 images in the test set.

The goal of this example is to build a classifier that can correctly identify the digit in an image. We will walk through the steps of loading the data, training a model, and evaluating the model.


```python
# We need to import the necessary libraries
# numpy as np
...
# matplotlib.pyplot as plt
...
# sklearn.datasets import fetch_openml
...
# sklearn.model_selection import train_test_split
...
# sklearn.linear_model import LogisticRegression
...
# sklearn.preprocessing import StandardScaler
...
# Load the MNIST dataset
# sklearn.datasets.fetch_openml is used to load the MNIST dataset, use the key 'mnist_784', and version 1, version is an optional parameter so set is as (version=1)
# assign the result to the variable mnist
...
# mnist.data is the data, mnist.target is the labels, assign these to X and y respectively
...
# Split the data into training and test sets
# sklearn.model_selection.train_test_split is used to split the data into training and test sets, test_size
# make the the proportion of the dataset to include in the test split to be 20% of the data, random_state is the seed used by the random number generator, set to 42 for reproducibility and testing - Note 42 is the answer to the universe
# random_state is the seed used by the random number generator, set to 42 for reproducibility and testing - Note 42 is the answer to the universe
...

# Standardize the data
# sklearn.preprocessing.StandardScaler is used to standardize the data, fit the scaler to the training data and then transform the training and test data
# save the scaler as an object scaler
...
# fit the scaler to the training data and then transform the training and test data, assign the transformed data to X_train_scaled and X_test_scaled respectively
# you can use the fit_transform method to fit the scaler to the training data and then transform the training and test data
# fit the scaler to the training data and then transform the training and test data, assign the transformed data to X_train_scaled and X_test_scaled respectively
...
# Train a Logistic Regression model
# use the LogisticRegression class to train a model. and save the object to the variable clf which stands for classifier
# set the max_iter to 1000, solver to 'lbfgs', and random_state to 42 for reproducibility and testing
...
# fit the model to the training data, use the fit method to fit the model to the training data - since this is an object we use the dot notation to call the fit method, and we do not need to assign the result to a variable
...
# Predict on the test set
# use the predict method to predict the labels of the test data, assign the result to the variable y_pred
...
# Note, we are actually going to train an machine learning model. It might take a minute or two to train - on my computer it takes about 30 seconds.
```

```python
# We have provided some code to plot and display a random image from the test set, and print the true and predicted labels.
# Did this do a good job?

# Choose a random test sample
random_index = np.random.randint(0, len(X_test))
random_image = X_test.iloc[random_index].values.reshape(28, 28)
true_label = y_test.iloc[random_index]
predicted_label = y_pred[random_index]

# Plot the image with prediction and true label
plt.figure(figsize=(4, 4))
plt.imshow(random_image, cmap="gray")
plt.title(f"True Label: {true_label}\nPredicted: {predicted_label}", fontsize=14)
plt.axis("off")
plt.show()
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week9-readings", "10_scikit_learn_q")
```

Tests for mnist-classifier (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> import matplotlib.pyplot as plt
>>> from sklearn.datasets import fetch_openml
>>> from sklearn.model_selection import train_test_split
>>> from sklearn.linear_model import LogisticRegression
>>> from sklearn.preprocessing import StandardScaler
>>> from sklearn.metrics import accuracy_score
>>> max_question_points = str(18.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(18.0)
>>> log_variable('total-points', f'week9-readings, 10_scikit_learn_q', 18.0)
>>> question_id = 'mnist-classifier-1'
>>> max_score = 2.0
>>> score = 0
>>> condition1 = hasattr(mnist, 'data') and hasattr(mnist, 'target')
>>> condition2 = mnist is not None
>>> assert condition1 and condition2, 'MNIST dataset does not have expected attributes'
>>> if condition1 and condition2:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_scikit_learn_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: MNIST dataset failed to load.
success_message: Success: MNIST dataset loaded successfully.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> import matplotlib.pyplot as plt
>>> from sklearn.datasets import fetch_openml
>>> from sklearn.model_selection import train_test_split
>>> from sklearn.linear_model import LogisticRegression
>>> from sklearn.preprocessing import StandardScaler
>>> from sklearn.metrics import accuracy_score
>>> question_id = 'mnist-classifier-2'
>>> max_score = 2.0
>>> score = 0
>>> X, y = (mnist.data, mnist.target)
>>> condition1 = isinstance(X, np.ndarray) or hasattr(X, 'shape')
>>> condition2 = isinstance(y, np.ndarray) or hasattr(y, 'shape')
>>> assert condition1 and condition2, 'X and y are not valid numpy arrays or dataframes'
>>> assert len(X) == len(y), 'Mismatch in data and target length'
>>> if (condition1 and condition2) and len(X) == len(y):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_scikit_learn_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Data and labels assignment failed.
success_message: Success: Data and labels assigned correctly.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> import matplotlib.pyplot as plt
>>> from sklearn.datasets import fetch_openml
>>> from sklearn.model_selection import train_test_split
>>> from sklearn.linear_model import LogisticRegression
>>> from sklearn.preprocessing import StandardScaler
>>> from sklearn.metrics import accuracy_score
>>> question_id = 'mnist-classifier-3'
>>> max_score = 2.0
>>> score = 0
>>> X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.2, random_state=42)
>>> assert len(X_train) == 56000 and len(X_test) == 14000, 'Incorrect train-test split sizes'
>>> assert len(y_train) == 56000 and len(y_test) == 14000, 'Mismatch in train-test target sizes'
>>> if (len(X_train) == 56000 and len(X_test) == 14000) and (len(y_train) == 56000 and len(y_test) == 14000):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_scikit_learn_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(len(X_train)))
>>> responses = update_responses(question_id, str(len(X_test)))
```
failure_message: Train-test split failed.
success_message: Success: Train-test split performed correctly.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> import matplotlib.pyplot as plt
>>> from sklearn.datasets import fetch_openml
>>> from sklearn.model_selection import train_test_split
>>> from sklearn.linear_model import LogisticRegression
>>> from sklearn.preprocessing import StandardScaler
>>> from sklearn.metrics import accuracy_score
>>> question_id = 'mnist-classifier-4'
>>> max_score = 2.0
>>> score = 0
>>> scaler = StandardScaler()
>>> X_train_scaled = scaler.fit_transform(X_train)
>>> X_test_scaled = scaler.transform(X_test)
>>> assert np.isclose(np.mean(X_train_scaled), 0, atol=0.1), 'Training data is not standardized correctly'
>>> assert np.isclose(np.std(X_train_scaled), 1, atol=0.1), 'Training data scaling is incorrect'
>>> assert np.isclose(np.mean(X_test_scaled), 0, atol=0.1), 'Test data is not standardized correctly'
>>> if np.isclose(np.mean(X_train_scaled), 0, atol=0.1) and np.isclose(np.std(X_train_scaled), 1, atol=0.1) and np.isclose(np.mean(X_test_scaled), 0, atol=0.1):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_scikit_learn_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(np.mean(X_train_scaled)))
>>> responses = update_responses(question_id, str(np.std(X_train_scaled)))
```
failure_message: StandardScaler transformation failed.
success_message: Success: StandardScaler initialized and transformed correctly.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> import matplotlib.pyplot as plt
>>> from sklearn.datasets import fetch_openml
>>> from sklearn.model_selection import train_test_split
>>> from sklearn.linear_model import LogisticRegression
>>> from sklearn.preprocessing import StandardScaler
>>> from sklearn.metrics import accuracy_score
>>> question_id = 'mnist-classifier-5'
>>> max_score = 2.0
>>> score = 0
>>> clf = LogisticRegression(max_iter=1000, solver='lbfgs', random_state=42)
>>> out = clf.fit(X_train_scaled, y_train)
>>> condition1 = hasattr(clf, 'coef_')
>>> assert condition1, 'Model coefficients are missing, model may not have been trained properly'
>>> if condition1:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_scikit_learn_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: LogisticRegression model training failed.
success_message: Success: LogisticRegression model trained successfully.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> import matplotlib.pyplot as plt
>>> from sklearn.datasets import fetch_openml
>>> from sklearn.model_selection import train_test_split
>>> from sklearn.linear_model import LogisticRegression
>>> from sklearn.preprocessing import StandardScaler
>>> from sklearn.metrics import accuracy_score
>>> question_id = 'mnist-classifier-6'
>>> max_score = 2.0
>>> score = 0
>>> y_pred = clf.predict(X_test_scaled)
>>> assert len(y_pred) == len(y_test), 'Mismatch in predicted and actual test labels'
>>> if len(y_pred) == len(y_test):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_scikit_learn_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(len(y_pred)))
```
failure_message: Prediction generation failed.
success_message: Success: Predictions generated successfully.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> import matplotlib.pyplot as plt
>>> from sklearn.datasets import fetch_openml
>>> from sklearn.model_selection import train_test_split
>>> from sklearn.linear_model import LogisticRegression
>>> from sklearn.preprocessing import StandardScaler
>>> from sklearn.metrics import accuracy_score
>>> question_id = 'mnist-classifier-7'
>>> max_score = 6.0
>>> score = 0
>>> accuracy = accuracy_score(y_test, y_pred)
>>> assert accuracy > 0.85, f'Model accuracy too low: {accuracy:.2f}'
>>> if accuracy > 0.85:
...     score = 6.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_scikit_learn_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Model accuracy too low.
success_message: Success: Model accuracy is sufficient.
