```python
from pykubegrader.tokens.validate_token import validate_token
validate_token('type the key provided by your instructor here', assignment = 'week10-lecture')

```

```python
validate_token(assignment = 'week10-lecture')

# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("12_attendance", "week_10", "lecture", assignment_points = 25.0, assignment_tag = 'week10-lecture')

# Initialize Otter
import otter
grader = otter.Notebook("12_attendance.ipynb")
```

# 🗓️📚 Week 10: Attendance ✅

![](./assets/figures/graduated-cb-shootn.gif)

To get credit for attendance make sure you run all blocks of code. 

```python
I_attended_lecture = True
```

```python
grader.check("week-10-lecture-attendance")
```

## 🎉 Well Done!

Keep up the good work.

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.tokens.validate_token import validate_token
validate_token(assignment = 'week10-lecture')


from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week10-lecture", "12_attendance")
```

Tests for week-10-lecture-attendance (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import matplotlib
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week10-lecture')
>>> import matplotlib
>>> matplotlib.use('Agg')
>>> max_question_points = str(25.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(25.0)
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week10-lecture')
>>> log_variable('total-points', f'week10-lecture, 12_attendance', 25.0)
>>> question_id = 'week-10-lecture-attendance-1'
>>> max_score = 25.0
>>> score = 0
>>> assert I_attended_lecture, 'You must attend the lecture to get credit for this assignment.'
>>> if I_attended_lecture:
...     score = 25.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('12_attendance', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Output does not include 'I attended the lecture'.
success_message: Correctly included the phrase 'I attended the lecture'.
