```python
from pykubegrader.tokens.validate_token import validate_token
validate_token('type the key provided by your instructor here', assignment = 'week9-lecture')

```

```python
validate_token(assignment = 'week9-lecture')

# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("4_clashroyal", "week_9", "lecture", assignment_points = 73.0, assignment_tag = 'week9-lecture')

# Initialize Otter
import otter
grader = otter.Notebook("4_clashroyal.ipynb")
```

# Clash Royal Object Detection

In this guideded tutorial we will walk through the process of detecting objects in an image using the `scikit-image` library.

We will start by loading the image and converting it to grayscale.


## Visualizing the images

We will start by visualizing the two images that we will be using for this tutorial. The first image is a background image and the second image a battle scene between Professor Agar and a random Foe.

Having the background image allows us to detect objects that are not part of the background.

```python
# Import necessary libraries
# Scikit-image is a Python library for image processing. We will use the io, color, filters, feature, measure, and morphology modules.
# from scipy import ndimage as ndi
# matplotlib and numpy using common aliases plt and np
...

# Load the Clash Royale image
# use the io module to load the image, from the assets/figures directory
# the background image is clash1.jpeg and the image to detect is clash2.jpeg
...


# If the image has 4 channels (RGBA), convert it to RGB by dropping the alpha channel
# hint, you can use the [..., :3] syntax to drop the alpha channel
...

# Visualize the images
# build a figure with two subplots, one for the background and one for the image
# use the figsize parameter to set the size of the figure to 10x5
# use the axis parameter to turn off the axis labels
# use the imshow function to display the images
# use the title function to set the title of the subplots, the title of the first image is "Background" and the title of the second image is "Image"
# use the show function to display the figure
...

plt.show()
```

```python
grader.check("Clash-Royal-Object-Detection-loading-images")
```

```python
# Convert the images to grayscale
# use the color module method rgb2gray to convert the images to grayscale, assign the results to image_gray and background_gray
...

# Apply edge detection filters
# use the filters module method sobel to apply the sobel edge detection filter to the background_gray and image_gray images, assign the results to edges_sobel_background and edges_sobel_image
# use the filters module method scharr to apply the scharr edge detection filter to the background_gray and image_gray images, assign the results to edges_scharr_background and edges_scharr_image
# use the feature module method canny to apply the canny edge detection filter to the background_gray and image_gray images, assign the results to edges_canny_background and edges_canny_image
...

# Combine edge filters for improved segmentation
# use the numpy module method add to combine the edges_sobel_background, edges_scharr_background, and edges_canny_background images, assign the result to combined_edges_background
# use the numpy module method add to combine the edges_sobel_image, edges_scharr_image, and edges_canny_image images, assign the result to combined_edges_image
...

# Visualize the combined edge filters
# use the matplotlib module method figure to create a figure with two subplots, one for the background and one for the image, make the figure size 10x5
# use the matplotlib module method subplot to create the subplots, the first subplot is for the background and the second is for the image
# use the matplotlib module method imshow to display the combined_edges_background and combined_edges_image images
# use the matplotlib module method title to set the title of the subplots, the title of the first subplot is "Background" and the title of the second subplot is "Image"
# use the matplotlib module method show to display the figure
fig, ax = plt.subplots(1, 2, figsize=(10, 5))

ax[0].imshow(combined_edges_background, cmap="gray")
ax[0].set_title("Background")
ax[0].axis("off")

ax[1].imshow(combined_edges_image, cmap="gray")
ax[1].set_title("Image")
ax[1].axis("off")

plt.show()
```

```python
grader.check("Edge-Detection-and-Thresholding")
```

```python
# Compute the difference between the combined edge detection for the background and image
# use the numpy module method abs to compute the absolute difference between the combined_edges_background and combined_edges_image arrays, assign the result to edge_diff
...

fig, ax = plt.subplots(figsize=(10, 5))
ax.imshow(edge_diff, cmap="gray")
ax.set_title("Edge Difference")
ax.axis("off")
plt.show()
```

```python
grader.check("edge-difference")
```

```python
# Thresholding to create a binary mask, where the threshold is 0.8
...

fig, ax = plt.subplots(1, 1, figsize=(10, 5))
ax.imshow(binary_mask, cmap="gray")
ax.set_title("Binary Mask")
ax.axis("off")
plt.show()
```

```python
grader.check("binary-mask")
```

```python
# Perform morphological closing to close small gaps in the edges
# use the morphology module method closing to perform morphological closing on the binary_mask image, assign the result to closed_mask
# use the morphology module method square to create a square structuring element with a radius of 5
...

fig, ax = plt.subplots(1, 1, figsize=(10, 5))
ax.imshow(closed_mask, cmap="gray")
ax.set_title("Closed Mask")
ax.axis("off")
plt.show()

plt.show()
```

```python
grader.check("closing-mask")
```

```python
# Merge connected shapes using morphological dilation
# use the morphology module method dilation to perform morphological dilation on the closed_mask image, assign the result to merged_mask
# use the morphology module method square to create a square structuring element with a radius of 10
...

# Visualize the merged mask
fig, ax = plt.subplots(1, 1, figsize=(10, 5))
ax.imshow(merged_mask, cmap="gray")
ax.set_title("Merged Mask")
ax.axis("off")
plt.show()
```

```python
grader.check("meger-mask")
```

```python
# Remove small objects (noise) from the binary mask
# use the morphology module method remove_small_objects to remove small objects from the merged_mask image, assign the result to cleaned_mask
# use the min_size parameter to set the minimum size to 5000, this is in pixels.
...

fig, ax = plt.subplots(1, 1, figsize=(10, 5))
ax.imshow(cleaned_mask, cmap="gray")
ax.set_title("Cleaned Mask")
ax.axis("off")
plt.show()

plt.show()
```

```python
grader.check("clean-mask")
```

```python
# from skimage import io, measure, color
# from skimage.measure import regionprops
...

# Threshold the image to create a binary mask
# create a binary mask from the cleaned_mask image, where the threshold is 0.8, assign the result to binary_mask
...

# Label connected regions in the binary image
# use the measure module method label to label the connected regions in the binary_mask image, assign the result to labels
# use the connectivity parameter to set the connectivity to 2
...

# Get properties of labeled regions
# use the regionprops function to get the properties of the labeled regions in the labels image, assign the result to regions
...

# Plot the segmented image
# use the matplotlib module method figure to create a figure with a single subplot, make the figure size 10x10
# use the matplotlib module method subplot to create the subplot
# use the matplotlib module method imshow to display the color.label2rgb(labels, bg_label=0) image
# use the matplotlib module method title to set the title of the subplot, the title is "Segmented Image with Bounding Boxes"
# use the matplotlib module method axis to turn off the axis labels
# use the matplotlib module method show to display the figure
...

# Draw bounding boxes around each labeled region
# use a for loop to iterate over the regions
# use the region.bbox attribute to get the coordinates of the bounding box
# use the matplotlib module method Rectangle to draw the bounding box on the image
# use the matplotlib module method add_patch to add the bounding box to the image
# use the matplotlib module method show to display the figure
...
    # Get the coordinates of the bounding box
    # use the region.bbox attribute to get the coordinates of the bounding box, assign the result to bbox
    # use the region.bbox attribute to get the coordinates of the bounding box, assign the result to bbox assign the results to min_row, min_col, max_row, max_col
    ...

    # Draw the rectangle on the image
    # use the matplotlib module method Rectangle to draw the bounding box on the image, assign the result to rect
    # use the edgecolor parameter to set the edge color to red
    # use the facecolor parameter to set the face color to none
    # use the linewidth parameter to set the line width to 2
    # the rectangle takes the coordinates (min_col, min_row) and the width and height of the rectangle (max_col - min_col, max_row - min_row)
    ...
    
    # use the matplotlib module method add_patch to add the bounding box to the image
    # use the ax.add_patch method to add the bounding box to the image
    ...

plt.show()
```

```python
grader.check("segment-image-regions")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.tokens.validate_token import validate_token
validate_token(assignment = 'week9-lecture')


from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week9-lecture", "4_clashroyal")
```

Tests for Clash-Royal-Object-Detection-loading-images (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week9-lecture')
>>> import sys
>>> max_question_points = str(11.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(73.0)
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week9-lecture')
>>> log_variable('total-points', f'week9-lecture, 4_clashroyal', 73.0)
>>> question_id = 'Clash-Royal-Object-Detection-loading-images-1'
>>> max_score = 1.0
>>> score = 0
>>> assert 'skimage' in sys.modules, 'scikit-image is not imported.'
>>> assert 'scipy' in sys.modules, 'scipy is not imported.'
>>> assert 'matplotlib.pyplot' in sys.modules, 'matplotlib.pyplot is not imported.'
>>> assert 'numpy' in sys.modules, 'numpy is not imported.'
>>> if 'skimage' in sys.modules and 'scipy' in sys.modules and ('matplotlib.pyplot' in sys.modules) and ('numpy' in sys.modules):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_clashroyal', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: Required libraries are not imported correctly.
success_message: Success: Required libraries are imported correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week9-lecture')
>>> question_id = 'Clash-Royal-Object-Detection-loading-images-2'
>>> max_score = 1.0
>>> score = 0
>>> assert 'background' in globals(), 'Background image is not defined.'
>>> assert isinstance(background, np.ndarray), 'Background is not a numpy array.'
>>> assert background.size > 0, 'Background image is empty.'
>>> if 'background' in globals() and isinstance(background, np.ndarray) and (background.size > 0):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_clashroyal', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(background))
```
failure_message: Failed: Background image is not loaded correctly.
success_message: Success: Background image is loaded correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week9-lecture')
>>> question_id = 'Clash-Royal-Object-Detection-loading-images-3'
>>> max_score = 1.0
>>> score = 0
>>> assert 'image' in globals(), 'Image to detect is not defined.'
>>> assert isinstance(image, np.ndarray), 'Image to detect is not a numpy array.'
>>> assert image.size > 0, 'Image to detect is empty.'
>>> if 'image' in globals() and isinstance(image, np.ndarray) and (image.size > 0):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_clashroyal', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(image))
```
failure_message: Failed: Image to detect is not loaded correctly.
success_message: Success: Image to detect is loaded correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week9-lecture')
>>> question_id = 'Clash-Royal-Object-Detection-loading-images-4'
>>> max_score = 4.0
>>> score = 0
>>> background_color_sum = np.sum(background)
>>> expected_background_sum = 1194260739
>>> assert background_color_sum == expected_background_sum, 'Incorrect background image loaded.'
>>> if background_color_sum == expected_background_sum:
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_clashroyal', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(background))
```
failure_message: Failed: Incorrect background image loaded based on color sum.
success_message: Success: Correct background image loaded based on color sum!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week9-lecture')
>>> question_id = 'Clash-Royal-Object-Detection-loading-images-5'
>>> max_score = 4.0
>>> score = 0
>>> image_color_sum = np.sum(image)
>>> expected_image_sum = 1175362976
>>> assert image_color_sum == expected_image_sum, 'Incorrect image to detect loaded.'
>>> if image_color_sum == expected_image_sum:
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_clashroyal', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(image))
```
failure_message: Failed: Incorrect image to detect loaded based on color sum.
success_message: Success: Correct image to detect loaded based on color sum!


Tests for Edge-Detection-and-Thresholding (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week9-lecture')
>>> max_question_points = str(20.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(73.0)
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week9-lecture')
>>> log_variable('total-points', f'week9-lecture, 4_clashroyal', 73.0)
>>> question_id = 'Edge-Detection-and-Thresholding-1'
>>> max_score = 2.0
>>> score = 0
>>> test = np.sum(edges_sobel_background)
>>> condition = np.isclose(test, 82125.810787)
>>> assert condition
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_clashroyal', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: sobel edge detection is not applied correctly.
success_message: Success: sobel edge detection is applied correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week9-lecture')
>>> question_id = 'Edge-Detection-and-Thresholding-2'
>>> max_score = 2.0
>>> score = 0
>>> test = np.sum(edges_canny_background)
>>> condition = np.isclose(test, 129762)
>>> assert condition
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_clashroyal', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: canny edge detection is not applied correctly.
success_message: Success: canny edge detection is applied correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week9-lecture')
>>> question_id = 'Edge-Detection-and-Thresholding-3'
>>> max_score = 2.0
>>> score = 0
>>> test = np.sum(edges_scharr_background)
>>> condition = np.isclose(test, 82690.87100435414)
>>> assert condition
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_clashroyal', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: canny edge detection is not applied correctly.
success_message: Success: canny edge detection is applied correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week9-lecture')
>>> question_id = 'Edge-Detection-and-Thresholding-4'
>>> max_score = 2.0
>>> score = 0
>>> test = np.sum(edges_sobel_image)
>>> condition = np.isclose(test, 82893.86893900018)
>>> assert condition
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_clashroyal', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: sobel edge detection is not applied correctly.
success_message: Success: sobel edge detection is applied correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week9-lecture')
>>> question_id = 'Edge-Detection-and-Thresholding-5'
>>> max_score = 2.0
>>> score = 0
>>> test = np.sum(edges_canny_image)
>>> condition = np.isclose(test, 132478)
>>> assert condition
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_clashroyal', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: canny edge detection is not applied correctly.
success_message: Success: canny edge detection is applied correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week9-lecture')
>>> question_id = 'Edge-Detection-and-Thresholding-6'
>>> max_score = 2.0
>>> score = 0
>>> test = np.sum(edges_scharr_image)
>>> condition = np.isclose(test, 83497.336501897)
>>> assert condition
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_clashroyal', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: canny edge detection is not applied correctly.
success_message: Success: canny edge detection is applied correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week9-lecture')
>>> question_id = 'Edge-Detection-and-Thresholding-7'
>>> max_score = 4.0
>>> score = 0
>>> test = np.sum(combined_edges_background)
>>> condition = np.isclose(test, 294578.6817915886)
>>> assert condition
>>> if condition:
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_clashroyal', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: combined edge detection is not applied correctly.
success_message: Success: combined edge detection for background is applied correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week9-lecture')
>>> question_id = 'Edge-Detection-and-Thresholding-8'
>>> max_score = 4.0
>>> score = 0
>>> test = np.sum(combined_edges_image)
>>> condition = np.isclose(test, 298869.2054408976)
>>> assert condition
>>> if condition:
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_clashroyal', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: combined edge detection for image is not applied correctly.
success_message: Success: combined edge detection for image is applied correctly!


Tests for binary-mask (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week9-lecture')
>>> max_question_points = str(4.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(73.0)
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week9-lecture')
>>> log_variable('total-points', f'week9-lecture, 4_clashroyal', 73.0)
>>> question_id = 'binary-mask-1'
>>> max_score = 4.0
>>> score = 0
>>> test = np.sum(binary_mask)
>>> condition = np.isclose(test, 62446, atol=1e-06)
>>> assert condition
>>> if condition:
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_clashroyal', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: binary mask is not applied correctly.
success_message: Success: binary mask is applied correctly!


Tests for clean-mask (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week9-lecture')
>>> max_question_points = str(4.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(73.0)
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week9-lecture')
>>> log_variable('total-points', f'week9-lecture, 4_clashroyal', 73.0)
>>> question_id = 'clean-mask-1'
>>> max_score = 4.0
>>> score = 0
>>> test = np.sum(cleaned_mask)
>>> condition = np.isclose(test, 370431)
>>> assert condition
>>> if condition:
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_clashroyal', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: combined edge detection is not applied correctly.
success_message: Success: combined edge detection for background is applied correctly!


Tests for closing-mask (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week9-lecture')
>>> max_question_points = str(4.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(73.0)
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week9-lecture')
>>> log_variable('total-points', f'week9-lecture, 4_clashroyal', 73.0)
>>> question_id = 'closing-mask-1'
>>> max_score = 4.0
>>> score = 0
>>> test = np.sum(closed_mask)
>>> condition = np.isclose(test, 121935)
>>> assert condition
>>> if condition:
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_clashroyal', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: closing mask is not applied correctly.
success_message: Success: closing mask is applied correctly!


Tests for edge-difference (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week9-lecture')
>>> max_question_points = str(4.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(73.0)
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week9-lecture')
>>> log_variable('total-points', f'week9-lecture, 4_clashroyal', 73.0)
>>> question_id = 'edge-difference-1'
>>> max_score = 4.0
>>> score = 0
>>> test = np.sum(edge_diff)
>>> condition = np.isclose(test, 111380.53573141411, atol=1e-06)
>>> assert condition
>>> if condition:
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_clashroyal', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: edge difference is not applied correctly.
success_message: Success: edge difference is applied correctly!


Tests for meger-mask (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week9-lecture')
>>> max_question_points = str(4.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(73.0)
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week9-lecture')
>>> log_variable('total-points', f'week9-lecture, 4_clashroyal', 73.0)
>>> question_id = 'meger-mask-1'
>>> max_score = 4.0
>>> score = 0
>>> test = np.sum(merged_mask)
>>> condition = np.isclose(test, 482699)
>>> assert condition
>>> if condition:
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_clashroyal', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: merged mask is not applied correctly.
success_message: Success: merged mask is applied correctly!


Tests for segment-image-regions (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week9-lecture')
>>> import sys
>>> max_question_points = str(22.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(73.0)
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week9-lecture')
>>> log_variable('total-points', f'week9-lecture, 4_clashroyal', 73.0)
>>> question_id = 'segment-image-regions-1'
>>> max_score = 3.0
>>> score = 0
>>> condition1 = 'matplotlib.pyplot' in sys.modules
>>> condition2 = 'skimage' in sys.modules
>>> condition5 = 'regionprops' in sys.modules
>>> assert condition1, 'matplotlib.pyplot is not imported.'
>>> assert condition2, 'scikit-image is not imported.'
>>> if condition1 and condition2:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_clashroyal', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: Required libraries are not imported correctly.
success_message: Success: Required libraries are imported correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week9-lecture')
>>> question_id = 'segment-image-regions-2'
>>> max_score = 2.0
>>> score = 0
>>> assert isinstance(binary_mask, np.ndarray), 'binary_mask is not a numpy array.'
>>> assert binary_mask.dtype == bool, 'binary_mask is not boolean.'
>>> if isinstance(binary_mask, np.ndarray) and binary_mask.dtype == bool:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_clashroyal', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(binary_mask))
```
failure_message: Failed: Binary mask is not created correctly.
success_message: Success: Binary mask is created correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week9-lecture')
>>> question_id = 'segment-image-regions-3'
>>> max_score = 2.0
>>> score = 0
>>> assert isinstance(labels, np.ndarray), 'labels is not a numpy array.'
>>> assert labels.shape == binary_mask.shape, 'labels do not match binary_mask shape.'
>>> if isinstance(labels, np.ndarray) and labels.shape == binary_mask.shape:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_clashroyal', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(labels))
```
failure_message: Failed: Labels are not created correctly.
success_message: Success: Labels are created correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week9-lecture')
>>> question_id = 'segment-image-regions-4'
>>> max_score = 3.0
>>> score = 0
>>> assert isinstance(regions, list), 'regions is not a list.'
>>> condition = all((hasattr(region, 'bbox') for region in regions))
>>> assert condition, 'Some regions do not have bbox attribute.'
>>> if isinstance(regions, list) and condition:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_clashroyal', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: Region properties are not computed correctly.
success_message: Success: Region properties are computed correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week9-lecture')
>>> question_id = 'segment-image-regions-5'
>>> max_score = 4.0
>>> score = 0
>>> computed_area_sum = sum((region.area for region in regions))
>>> expected_area_sum = 370431.0
>>> assert computed_area_sum == expected_area_sum
>>> if computed_area_sum == expected_area_sum:
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_clashroyal', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: Sum of region areas is incorrect.
success_message: Success: Sum of region areas is correct!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week9-lecture')
>>> question_id = 'segment-image-regions-6'
>>> max_score = 4.0
>>> score = 0
>>> computed_perimeter_sum = sum((region.perimeter for region in regions))
>>> expected_perimeter_sum = 31773.5370747862
>>> condition = np.isclose(computed_perimeter_sum, expected_perimeter_sum, atol=1e-06)
>>> assert condition
>>> if condition:
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_clashroyal', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: Sum of region perimeters is incorrect.
success_message: Success: Sum of region perimeters is correct!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week9-lecture')
>>> question_id = 'segment-image-regions-7'
>>> max_score = 4.0
>>> score = 0
>>> computed_eccentricity_sum = sum((region.eccentricity for region in regions))
>>> expected_eccentricity_sum = 13.643185559487907
>>> condition = np.isclose(computed_eccentricity_sum, expected_eccentricity_sum, atol=1e-06)
>>> assert condition
>>> if condition:
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_clashroyal', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: Sum of region eccentricities is incorrect.
success_message: Success: Sum of region eccentricities is correct!
